# Lab 1

This is the first lab of Data 100 and is intended to ensure that you have your Jupyter Hub and/or local Notebook environments configured correctly.

This assignment should be completed before Tuesday 23:00 at 1:00AM.


# Imports

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
%matplotlib inline
```

# Exercices

## Part A

Write a function that returns a list of numbers, such that $ x_i=i^2 $,for $1\leq i \leq n$. Don't worry about the case where $n < 0$.

<!-- grade_id: cell-9388c587d0e639ff, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """
    Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    """
    ### BEGIN SOLUTION
    if n < 1:
        raise ValueError("n must be greater than or equal to 1")
    return [i ** 2 for i in range(1, n + 1)]
    ### END SOLUTION
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$.  
Check that it does:

```python
squares(10)
```

Check that squares returns the correct output for several inputs.

<!-- grade_id: squares-tests, points: 1, grade: true, solution: false -->
```python
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
### BEGIN HIDDEN TESTS
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
### END HIDDEN TESTS
```

## Part B
Write a function which takes in a string and returns `True` if the string is a palindrome. (A string is a palindrome if it is the same forwards and backwards.)

<!-- grade_id: cell-b1fa3ce3a6566fbc, points: 0, grade: false, solution: true -->
```python
def is_palindrome(word):
    """Return True if word is a palindrome"""
    ### BEGIN SOLUTION
    middle = int(len(word) / 2)
    is_even = len(word) % 2 == 0
    left = word[:middle]
    right = word[:middle-is_even:-1]
    return left == right
    ### END SOLUTION
```

Your function should return true for "racecar".

```python
is_palindrome("racecar")
```

Check that the function works for several inputs.

<!-- grade_id: test-is-palindrome, points: 1, grade: true, solution: false -->
```python
assert is_palindrome("aviddiva") == True
assert is_palindrome("clearlynotapalindrome") == False
assert is_palindrome("kayak") == True
### BEGIN HIDDEN TESTS
assert is_palindrome("ab") == False
assert is_palindrome("abb") == False
assert is_palindrome("a") == True
### END HIDDEN TESTS
```

## Part C

Run the following code to see a pretty plot.

```python
from mpl_toolkits.mplot3d import axes3d
import matplotlib.pyplot as plt


# Create a plot
u = np.linspace(1.5*np.pi, -1.5*np.pi, 100)
[x,y] = np.meshgrid(u, u)
z = np.cos( np.sqrt(x.flatten()**2 + y.flatten()**2) ).reshape(x.shape)


fig = plt.figure()
ax = fig.add_subplot(111, projection='3d')
# Plot a basic wireframe.
ax.plot_wireframe(x, y, z, rstride=10, cstride=10)
ax.view_init(elev=50., azim=30)
plt.savefig("figure1.pdf")
```

# Part D
Do you think that eating french fries with mayonnaise is a crime?  
Tell us what you think in the following Markdown cell :)

<!-- grade_id: cell-ee0714a310fcdcf5, points: 0, grade: true, solution: true -->
